```python
# Initialize Otter
import otter
grader = otter.Notebook("lab12.worksheet.ipynb")
```

# Lab 12 - Worksheet 

Complete and turn in by Nov 17, 7pm.

```python
# Run this cell to load all required Python libraries 
import numpy as np
from datascience import *

import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline

import warnings
warnings.simplefilter('ignore')
```

# Question 1 

## Data

We load datasets in from other sources using the `Table.read_table()` method. We can pass in a *file path* to this method and it will load that data into a table we can use in Python!

The dataset that we'll use in this lab comes from the Behavioral Risk Factor Surveillance System (BRFSS), a health survey fielded by the Centers for Disease Control and Prevention (CDC). From the [BRFSS website](https://www.cdc.gov/brfss/index.html):
>The Behavioral Risk Factor Surveillance System (BRFSS) is the nation’s premier system of health-related telephone surveys that collect state data about U.S. residents regarding their health-related risk behaviors, chronic health conditions, and use of preventive services.

>By collecting behavioral health risk data at the state and local level, BRFSS has become a powerful tool for targeting and building health promotion activities. 

The dataset that you will investigate is a **subset of the 2022 BRFSS Survey**. We've taken all the data points corresponding to fully-completed surveys and, in our opinion, the most interesting columns. Since the entire data set is so large, we've randomly sampled 100,000 rows from the original data. While we've wrangled and cleaning the data set you'll use in your investigation, you're welcome to investigate the original source; you can do so via the [Survey Data section](https://www.cdc.gov/brfss/data_documentation/index.htm) of the BRFSS site.

```python
# Run this cell to load in the data
brfss = Table.read_table("data/brfss.csv")
brfss.show(5)
```

## Q1A 

Fill in the `num_rows_brfss` and `num_columns_brfss` with the number of rows and columns in the original `brfss` table, respectively.

```python
num_rows_brfss = ...
num_columns_brfss = ...
print(f"Our `brfss` table has {num_rows_brfss} rows and {num_columns_brfss} columns")
```

```python
grader.check("q1a")
```

### Investigating our Data

Now that we've loaded our data into the `brfss` table, let's take a closer look at its columns. Run the following cell to output the column names. 

```python
brfss.labels
```

Based on these column names, it looks like the data includes questions about **telecommunications**, **housing**, **demographic information**, **mental and physical health**, **alcohol and drug consumption**, and **physical exercise**. Each column in the `brfss` table corresponds to a question asked in the official BRFSS Survey.

## Q1B

Assign `num_alabama_rows` to the the number of times the name **Alabama** appeared in the `brfss` table.


```python
alabama_tbl = ...
num_alabama_rows = ...
num_alabama_rows
```

```python
grader.check("q1b")
```

Take a closer look at some of the columns in the `brfss` table. For the next two questions, we will be looking at the `"Binge Drinks"` column, which corresponds to this survey question:
> Considering all types of alcoholic beverages, how many times during the past 30 days did you have 5 or more
drinks for men or 4 or more drinks for women on an occasion? 

Notice that this column contains negative values, most notably `-1`. Why might this be the case? Discuss with the people around you and check in with your GSI to confirm.

```python
# Run this cell
brfss.column("Binge Drinks")
```

## Q1C 

Create a new table called `missing_binge_drinks` which only contains rows from the `brfss` table where there is a `-1` in the `"Binge Drinks"` column.

```python
missing_binge_drinks = ...
missing_binge_drinks
```

```python
grader.check("q1c")
```

<!-- BEGIN QUESTION -->

## Q1D 

Say we wanted to find the average of one of the columns from our original table. How does the inclusion of `-1` values *affect* this average? If we removed all the negative values, how would the average change?

*Fill in your answer here*

<!-- END QUESTION -->

Then, once you've answered, run the following cell to confirm your understanding.

```python
brfss_no_negatives_children = brfss.where("Children", are.not_equal_to(-1))
print(f"With negatives: {np.average(brfss.column('Children'))}")
print(f"Without negatives: {np.average(brfss_no_negatives_children.column('Children'))}")
```

## Q1E 

As an **_extra challenge_**, see if you can output Alabama's average number of children for all respondents who gave an answer (i.e. do not have a value of -1) **without using the `np.average` function**. You should use the table `brfss_no_negatives_children` that's defined in the cell above.

*Note*: We've provided some starter code. Feel free to use it or try another approach!

```python
alabama_tbl = ...
children_total = ...
average_num_children = ...
average_num_children
```

```python
grader.check("q1e")
```

## Q1F

Using the `no_missing_income` table we've provided for you, determine the **average income** (from `Income Upper`) for respondents who:

1. Have a health insurance plan ("Yes" in "Health Plan" column)
2. Do not have a health insurance plan ("No" in "Health Plan" column)
3. Refused to answer this question ("Declined to Answer" in "Health Plan" column)

```python
no_missing_income = brfss.where("Income Upper", are.not_equal_to(-1))

health_plan = ...
no_health_plan = ...
declined = ...

average_plan = ...
average_no_plan = ...
average_declined = ...
print(f"Respondents with a Health Insurance Plan made: \t\t${round(average_plan, 2)}")
print(f"Respondents without a Health Insurance Plan made: \t${round(average_no_plan, 2)}")
print(f"Respondents who refused to answer made: \t\t${round(average_declined, 2)}")
```

```python
grader.check("q1f")
```

# Question 2 

We will continue looking at the `brfss` data set, but also look at specific aspects of this data. 

### `barh`

The `barh` (horizontal bar chart) method is used to visualize **categorical** variable values. Categorical variables are non-numbers, like names and qualities (Color, State Names, etc.). As we saw in lecture, categorical variables come in 2 different types: *ordinal* and *nominal*. 

The `barh` method takes in 1 mandatory argument, which is the **name of the column** you want on the left (vertical) axis of your `barh` plot. There are also optional arguments that have to do with plotting -- you'll see examples of those in this lab and in the homework. The remaining optional arguments in the `datascience` documentation linked above can also be used, feel free to try out some of the others on your own!

To use the `barh` method properly, we first need to select the columns we want to see in the graph. We should not call `barh` directly on a large `Table` because without specifying a column, we get a bar graph for every single instance of every single variable, which you can imagine results in a lot of bar graphs.

<!-- BEGIN QUESTION -->

## Q2A 

Plot a horizontal bar chart that shows the counts of each category from the `"Days Smoking"` column of the `brfss` table. 

*Hint*: Use the `smoking_counts` table.

```python
smoking_counts = Table.read_table("data/smoking_counts.csv")
...
```

<!-- END QUESTION -->

### Multiple Columns 

We can also use `barh` to see multiple statistics at once. Let's use the `barh` method to see the average number of both *poor mental health* and *poor physical health* days. We'll be using the following columns:
1. `"Physical Health"`: Now thinking about your physical health, which includes physical illness and injury, for how many days during the past 30 days was your physical health not good? 
2. `"Mental Health"`: Now thinking about your mental health, which includes stress, depression, and problems with emotions, for how many days during the past 30 days was your mental health not good? 

Run the following cell to show an example of how to create an *overlaid bar chart* with two statistics.

```python
state_averages = Table.read_table("data/state_averages.csv")
state_averages
```

```python
# We must group first to get our desired columns, then we can call barh
state_averages.barh("State", overlay=True)
```

If we want different visualizations for each variable, we can set the optional `overlay` argument to `False`. The default value of `overlay` is `True`, so if you don't give it a value, you will get a plot with all the included variables at once.

```python
state_averages.barh("State", overlay=False)
```

That way we can choose if we want to have one plot with all our information or a new plot for each piece of information!

In this case, do we prefer an overlaid plot or two separate plots? Can you think of a case where we might want to have two separate plots instead of one overlaid plot? (Hint: think about the units for both variables — are they the same or different?)

Discuss with the people around you. 

### Where `barh` fails 

The `barh` method works well on categorical variables, but what if we have a **numerical** variable that we want to see the distribution in one particular state? Let's see what happens if we try to use `barh` on a numerical variable (`"Binge Drinking"`) instead of a categorical variable:

```python
# Just run this cell -- don't worry about this `group` method
brfss.group("Binge Drinks").barh("Binge Drinks")
```

As you can see, this bar plot is not particularly helpful. There are many categories that seem to not have any corresponding bar. Yet, that isn't the case! Seeing the breakdown of `"Binge Drinks"` does not provide us with any useful information, and it is also difficult to read or understand. Instead, for numerical variables, we have another visualization method that helps us visualize a numerical variable's distribution...

### `hist` 

The `hist` method allows us to see the distribution of a numerical variable. Categorical variables should be visualized using `barh`, and numerical variables should be visualized using `hist`.

The `hist` method takes in 1 mandatory argument and has several optional arguments (as is the case with `barh`, there are many other optional arguments, but here are just a few of them). For this lab, we'll set **`density` to be `False`**.

| **Argument** | **Description** | **Type** | **Mandatory?** |
| -- | -- | -- | -- |
| `column` | Column name whose values you want on the x-axis of your plot | Column name (string) | Yes |
| `density` | If `True`, then the resulting plot will be displayed not on the count of a value, but on the density of that value in the Table | boolean | No |
| `group` | Similar to the Table method `group`, groups rows by this label before plotting | Column name (string) | No |
| `overlay` | When `False`, make a new plot for each eligible statistic in the Table | boolean | No |
| `bins` | A NumPy array of bin boundaries you want your histogram to gather data into | array | No |
| `unit` |  A name for the units of the plotted column | Column name (string) | No |

**Again, in all cases, `density` should be set to `False`**

Keep in mind the same plotting optional arguments mentioned in the `barh` introduction.

Let's take a look at the distribution of exercise sessions in different states to see how the `hist` method helps visualize numerical variables, first starting with our favorite state, Michigan. We'll use the `sleep_no_negatives` table to exclude missing values (-1's).

```python
sleep_no_negatives = brfss.where("Sleep Time", are.not_equal_to(-1))
```

```python
# This plot shows the distribution of sleep time for Michiganders
my_bins = np.arange(0, 25, 1)
michigan = sleep_no_negatives.where("State", "Michigan")
michigan.hist("Sleep Time", density = False, bins=my_bins)
```

This shows us that people living in Michigan usually tend to sleep between 7 to 8 hours a night, but there are many people who sleep more hours (10+) or few hours (less than 5). Let's see how that compares to sleep time in another state, Indiana:

<!-- BEGIN QUESTION -->

## Q2B  

Fill in the following **code cell** to produce a histogram representing the ***distribution of sleep time*** for respondents from the state of Indiana.

*Note*: Set the optional `bins` argument of the `hist` method to `my_bins`. We've provided this variable for you.


```python
# This plot shows the distribution of sleep time for Indiana residents
my_bins = np.arange(0, 25, 1)
ind = ...
...
```

<!-- END QUESTION -->

### Michigan vs. Indiana



We can use `hist` on a `Table` with just rows for these two states and use the optional `group` argument.

*Note*: You'll see how `are.contained_in` works with the `where` method later on. For now, think of it as finding any rows corresponding to *either* `"Michigan"` or `"Indiana"`.

```python
# Just run this cell to load the `il_ca` table
in_mi = sleep_no_negatives.where("State", 
                                 are.contained_in(["Michigan", "Indiana"]))
in_mi.show(5)
```

<!-- BEGIN QUESTION -->

## Q2C 

Now that we've created our `il_mi` table, fill in the following **code cell** to produce a histogram representing the ***distribution of sleep time*** for *both* California and Illinois. You'll first need to `select` the necessary columns from `il_ca` then fill in the appropriate call to the `hist` method.

*Hint*: Take a look at the optional `group` argument from the description above.

*Note*: Set the optional `bins` argument of the `hist` method to `my_bins`. We've provided this variable for you.


```python
# plot shows the distribution of sleep time for people from Michigan AND Indiana
my_bins = np.arange(0, 25, 1)
...
```

<!-- END QUESTION -->

It appears that sleep time in Michian is a very similar, on average, to the sleep time in Indiana. The plot above shows the Michigan `Sleep Time` to be almost exactly on top of the Indiana's `Sleep Time`. Let's see if we can use a table query to figure out the same information:

```python
print(f"Michigan average:\t{np.mean(michigan.column('Sleep Time'))}")
print(f"Indiana average:\t{np.mean(ind.column('Sleep Time'))}")
```

As we can see, the plot we made appeared to suggest that average amount of sleep should be very similar between Michigan and Indiana, and the table operations reflected that! This is a benefit of visualization, that information can be learned about the dataset with just visual observation. It is always beneficial to back your claims about data with concrete facts about the dataset, but visualizations can help abstract away some of the confusion of looking at raw data so that non-data-scientists can better understand what is going on.

### Discussion 

Now, think about what would happen if you chose two states with **very different counts**, why would it be more difficult to compare them with histograms? 

Once you've discussed with someone around you, proceed with the code cells below to confirm your answers. We'll look to compare **Texas** and **Delaware**.


```python
# Just run this cell
texas = sleep_no_negatives.where("State", "Texas")
delaware = sleep_no_negatives.where("State", "Delaware")
print(f"Texans in `cleaned_exercise_mental_health` dataset: {texas.num_rows}")
print(f"Delawareans in `cleaned_exercise_mental_health` dataset: {delaware.num_rows}")
```

Each individual plot looks fine:

```python
# This plot shows the distribution of sleep times for Texas respondents
my_bins = np.arange(0, 25, 1)
texas.hist("Sleep Time", density = False, bins=my_bins)
```

```python
# This plot shows the distribution of sleep times for Delaware respondents
my_bins = np.arange(0, 25, 1)
delaware.hist("Sleep Time", density = False, bins=my_bins)
```

Take a look at the y-axis on both of these plots. What do you think will happen when we try to plot them on the same graph?

```python
# Just run this cell
texas_delaware = sleep_no_negatives.where("State", are.contained_in(["Texas", "Delaware"]))
texas_delaware.show(5)
```

<!-- BEGIN QUESTION -->

## Q2D 

Using the code in **Question 2C** as reference, produce a histogram showing the distribution of sleep time for respondents from *Delaware* and *Texas*. What do you notice about this plot?

```python
# plot shows the distribution of sleep times for people from Delaware and 
#  Texas exercised in a month
my_bins = np.arange(0, 25, 1)
...
```

<!-- END QUESTION -->

As you can see, there is so much more Texas data than Delaware data that we can hardly make comparisons between the two. Trying to figure out information from this plot is very difficult, so we would either have to use another type of visualization or change the perspective of this plot to be able to learn from it.

### `scatter` 

As we mentioned, visualizing two variables can show us patterns in the data that can help us learn new information. The `scatter` method allows us to see the relationship between two numerical variables in our data using a **scatter plot**. The first provided column name goes along the x-axis and the second goes along the y-axis.

Let's take a look at the relationship between **Physical Health** and **Alcohol Consumption**. For reference, here are the following questions from the original BRFSS Survey that correspond to our `"Physical Health"` and `"Binge Drinks"` columns.

> **Physical Health:** Now thinking about your physical health, which includes physical illness and injury, for how many days during
the past 30 days was your physical health not good? 


>**Binge Drinks**: Considering all types of alcoholic beverages, how many times during the past 30 days did you have 5 or more drinks for men or 4 or more drinks for women on an occasion?

## Q2E 

### Data Cleaning

As was the case with our previous visualizations lab, we know that the missing numerical values are encoded as `-1`s. Create a new table called `scatter_cleaned` which contains every row from the original `brfss` table that *does not* contain a `-1` in either the `"Physical Health"` column or the `"Binge Drinks"` column.

*Hint*: If you're having trouble with the code, feel free to reference the `barh` section of this lab.


```python

scatter_cleaned = ...
scatter_cleaned.show(5)
```

```python
grader.check("q2e")
```

### Scatter Plots 

Now, we can call `scatter` on the `scatter_cleaned` table. Run the following cell to do so.

```python
scatter_cleaned.scatter("Binge Drinks", "Physical Health")
```

Just like that, you've produced your first scatter plot! It looks a little messy, however. Oftentimes scatter plots can suffer from what's known as **[overplotting](https://www.displayr.com/what-is-overplotting/)**: when many data points fall on top of each other, creating a blob of data. When data is *overplot*, it's often difficult to see the individual data points on the scatter plot.

To fix this, we attempt to focus in on a smaller subset of the data. In this case, we'll look at points in which `"Binge Drinks"` falls between 0 and 30 days and the `"Physical Health"` column falls between 0 and 30 days.

```python
# Create a smaller subset of data
scatter_reduced = scatter_cleaned.where("Binge Drinks", 
                                are.below(30)).where("Physical Health", are.below(30))
scatter_reduced
```

<!-- BEGIN QUESTION -->

## Q2F

Using the `scatter_reduced` table, produce a scatterplot that plots "`Binge Drinks"` on the x-axis and `"Physcial Health"` on the y-axis. The code should be very similar to the previous scatter plot.

```python
...

```

<!-- END QUESTION -->

### The `group` and `labels` Optional Arguments

The `scatter` method also allows you to specify specific groups or labels for each data point using the `group` or `labels` keyword arguments. 

Say we wanted to investigate the relationship between an individual's **number of children** and their **mental health**. The corresponding questions from the original BRFSS survey were as follows:
>  **Mental Health**: Now thinking about your mental health, which includes stress, depression, and problems with emotions, for how many days during the past 30 days was your mental health not good?

> **Children**: How many children less than 18 years of age live in your household?



<!-- BEGIN QUESTION -->

## Q2G

In order to take advantage of the optional arguments, let's first load an additional table from the `"states_scatter.csv"` file. We'll provide the code for this.

Then, using the `states` table, produce a scatter plot that plots the average children against the average number of poor mental health days. 

```python
states = Table.read_table("data/states_scatter.csv")
states
```

```python
...

```

<!-- END QUESTION -->

This plot looks good, but it is difficult to see which points correspond to which states. To give each data point it's city name, we can use the `group` or `label` arguments:

```python
states.scatter("Children mean", "Mental Health mean", labels="State")
```

```python
states.scatter("Children mean", "Mental Health mean", group="State")
```

Neither of these plots are perfect given the number of data points, It is important to weigh the pros and cons of using each argument, as both will be more useful depending on the application

Scatter plots are useful when visualizing two numerical variables together. If you want to plot two numerical variables but one of those variables corresponds to time, we can use a line plot to visualize the non-time variable as time passes.

### `plot`

Similar to `scatter`, we give plot the names of two numerical columns and it creates a **line plot** for us. If we want to draw multiple line plots on the same set of axes, we give it a table with multiple numerical columns, and tell it which one contains the values for the x-axis.

The `plot` method allows us to see how non-time variables change over time. Let's use `plot` to look at the exercise patterns over the course of the year. First, we will look at a single line plot using `plot`:

```python
# Just run this cell to load a new table
months = Table.read_table("data/months.csv")
months
```

<!-- BEGIN QUESTION -->

## Q2H

Using the `months` table and the `plot` method, produce a *line plot* that plots the average sleep time over time

*Hint*: You'll want to plot the month on the x-is and average exercise sessions on the y-axis.


```python
...

```

<!-- END QUESTION -->

### Identifying Temporal Patterns
 
 Line plots are incredibly effective tools for identifying temporal patterns (i.e. changes over time). Let's utilize our newfound knowledge of the `plot` method to uncover underlying temporal patterns within our BRFSS data. Run the following cells and answer the question that follows.

```python
# Run this cell -- you should understand how this code works
vermont = sleep_no_negatives.where("State", "Vermont")
florida = sleep_no_negatives.where("State", "Florida")
```

```python
# Run this cell to produce a line plot for Vermont
vt_grouped = vermont.group("Month", np.average)
vt_grouped.plot("Month", "Sleep Time average")
```

```python
# Run this cell to produce a line plot for Florida
fl_grouped = florida.group("Month", np.average)
fl_grouped.plot("Month", "Sleep Time average")
```

<!-- BEGIN QUESTION -->

If we want to see multiple variables on one plot, we can include them in the table we call `plot` on. 

## Q2I 

For both the `data/vermont_averages.csv` and `data/florida_averages.csv` files, read the file into two new tables, `vt_health` and `fl_health`, respectively. Then, for each table, select the following columns:
1. Month
2. Physical Health average
3. Mental Health average

Finally, produce a scatter plot with *one line per variable* that is not `"Month"`. That is, `"Month"` is what should be plotted on the x-axis.

> ***Hint:** Check the top of the notebook for an example of how to load data* 


```python
vt_health = ...
fl_health = ...
```

```python
vt_health.show(5)
```

```python
# Create the line plot for Vermont 
...
```

```python
# Create the line plot for Florida 
...
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: 2)
```
>>> num_rows_brfss == 100000
True
```
```
>>> num_columns_brfss == 23
True
```


Tests for q1b (points: 1)
```
>>> num_alabama_rows == 1352
True
```


Tests for q1c (points: 2)
```
>>> missing_binge_drinks.num_rows == 51042
True
```


Tests for q1e (points: 2)
```
>>> np.isclose(average_num_children, 0.5443698732289336)
True
```


Tests for q1f (points: 3)
```
>>> np.isclose(average_plan,43832.35496992514)
True
```
```
>>> np.isclose(average_no_plan, 34940.17647058824)
True
```
```
>>> np.isclose(average_declined, 30963.912280701756)
True
```


Tests for q2e (points: 2)
```
>>> isinstance(scatter_cleaned, Table)
True
```
```
>>> scatter_cleaned.num_rows == 48403
True
```
