```python
# Initialize Otter
import otter
grader = otter.Notebook("lab11.worksheet.ipynb")
```

# Lab 11 - Worksheet 

Complete and turn in by Nov 2, 7pm. 

```python
from datascience import *
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
plt.style.use("ggplot")

```

# Part 1: Tables

## The Data

<img src = 'images/sfo.jpg' width = 600>

In this section, we'll look at airline data from San Francisco International Airport (SFO). The dataset was downloaded directly from [flysfo.com](https://data.sfgov.org/Transportation/Air-Traffic-Passenger-Statistics/rkru-6vcg), SFO's website. The dataset contains information about major airlines, their associated flight statistics, and other information related to their flights out of SFO.


<hr style="border: 1px solid #fdb515;" />


# Question 1 - Airlines


---
## Question 1a

Having gotten word of your newfound data science skills, the City of San Francisco tasks you with conducting data analysis on airline traffic out of San Francisco International Airport. To do so, they've given you a file containing relevant flight data. 

In the following cell, load the data contained in the `"data/sfo.csv"` file into the `sfo` table. You'll want to use one of the table methods reviewed in the last two weeks. 

[Python References](https://mtu.instructure.com/courses/1471027/pages/python-resources?module_item_id=22413830)

```python
# Read in data into a Table 

sfo ... 

sfo.show(10)
```

```python
grader.check("q1a")
```

---
## Question 1b

<img src="images/jetblue.png" width=300>

Now that you've loaded your data into the `sfo` table, it's time to conduct our analysis. The City of San Francisco wants you to first focus on **JetBlue Airways**.

Assign `num_passengers` to the **total number of passengers** on flights corresponding to **JetBlue Airways** in the `sfo` table. 

**Approach:** A recommended approach for solving this problem: 

1. create a new table, `jetblue`, containing *only* rows from the `sfo` table corresponding to `"JetBlue Airways"` as the `"Operating Airline"`.

2. access data in the `"Passenger Count"` column using the `.column` method, then calculate the **total number of passengers** using the `sum` method.  



```python
jetblue = ...   
num_passengers = ...
num_passengers
```

```python
grader.check("q1b")
```

---
## Question 1c

<img src="images/united.png" width=300>

Now, the City of San Francisco wants you to focus on **United Airlines**; they would like to know how many flights conducted by United Airrlines arrive in Europe. 

Assign `prop_europe` to the **proportion of flights** that **United Airlines** takes to Europe, as based on the information presented within the `sfo` table. 

**Approach**  A recommended approach for solving this problem: 

1. You'll first want to create a intermediary table, `united`, which contains data corresponding to `"United Airlines"` as the `"Operating Airline"`

2. Create a table from `united` with only those rows *where* the `"GEO Region"` is `"Europe"` and count the number of rows in this table.  
Divide this value by the number of rows for all **United Airlines** flights to get the proportion of flights.  

```python
united = ...   
prop_europe = ...
prop_europe
```

```python
grader.check("q1c")
```

---
## Question 1d

Let's continue our work with **United Airlines**, specifically looking at how **passenger counts** vary between successive years. In this question, we'll investigate two months: **July 2019** and **July 2020**.

Assign `july_2019_mean` to the **average number of passengers** on United Airlines flights in July 2019. Similarly, assign `july_2020_mean` to the **average number of passengers** on United Airlines flights in July 2020.

*Hint*: You can re-use the `united` table from the previous question.

**Approach**: A recommended approach for solving the problem is to create a new table *where* the "Activity Period" matches the time period of interest, select the column of interest that you want to do some math on, e.g., "Passenger Count", and take the average. 


```python
# Use the united Table from the prior question 

july_2019_mean = ...
july_2020_mean = ...


# Don't change any code below this line
print(f"Average Passenger Count (July 2019): {round(july_2019_mean, 2)}")
print(f"Average Passenger Count (July 2020): {round(july_2020_mean, 2)}")
```

```python
grader.check("q1d")
```

# Part 2 - Visualization 

*Acknowledgements*

Many of the pictures and descriptions in this section of the homework are taken from [Dr. Allison Horst](https://twitter.com/allison_horst) and [Dr. Kristen Gorman](https://www.uaf.edu/cfos/people/faculty/detail/kristen-gorman.php). The data was collected as a part of the [Long Term Ecological Research Network](https://lternet.edu/). If you're curious about the origins of the data, see [here](https://allisonhorst.github.io/palmerpenguins/) for more details.


## The Data

In this part of your assignment, we will explore a dataset containing size measurements for three penguin species observed on three islands in the Palmer Archipelago, Antarctica. The data was collected by Dr. Kristen Gorman, a marine biologist from 2007 to 2009.

Here's a photo of Dr. Gorman in the wild collecting the data:

<img src='images/gorman1.png' width=300>

Run the cell below to load in our data.

```python
# Run this cell!
penguins = Table.read_table("data/penguins.csv")
penguins
```

### Understanding the Data

Let's make sure we understand what each of the columns in our data represents before proceeding.

**`"species"`** There are three species of penguin in our dataset: Adelie, Chinstrap, and Gentoo.

<img src='images/lter_penguins.png' width=300>

**`"island"`** The penguins in our dataset come from three islands: Biscoe, Dream, and Torgersen. (The smaller image of Anvers Island may initially be confusing; the dark region is land and the light region is water.)

<img src='images/island.png' width=300>

Image taken from <a href="https://journals.plos.org/plosone/article/figure?id=10.1371/journal.pone.0090081.g001">here</a>


**`"bill_length_mm"`** and **`"bill_depth_mm"`**: See the illustration below.

<img src='images/culmen_depth.png' width=300>

**`"flipper_length_mm"`**: [Flippers](https://www.thespruce.com/flipper-definition-penguin-wings-385251) are the equivalent of wings on penguins. 

**`"body_mass_g"`** The body mass of the penguin, in grams

**`"sex"`**: The sex of the penguin

# Question 2 - Bar Charts

<!-- BEGIN QUESTION -->

---
## Question 2a
Let's start by visualizing the distribution of the islands from which the penguins in our dataset come. Your friend Sunya says that using a **histogram** is the best way to do this. However, Josh disagrees and says that using a **horizontal bar chart** is a better move. Who is right? Explain your answer.

*Hint*: What kind of variable is contained within the `"island"` column?

**State which plot is better and why?**

...

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---
## Question 2b

Based on your answer for the previous question, create an appropriate visualization to show the distribution of the islands from which the penguins come.

Specifically, write a line of code in the following cell that produces a visualization with three bars, one for each island. The length of each bar should correspond to the number of penguins from each island.

We've provided the `islands` table for you to use.

```python
islands = Table.read_table("data/islands.csv")
islands
```

```python

...

# Don't change any code below this line
plt.xlabel("Number of Penguins")
plt.title("Number of Penguins on Each Island");
```

<!-- END QUESTION -->

# Question 3 - Histograms


Great! Now that we've explored the distribution of a categorical variable in our dataset, it's time to study the distributions of some of the numerical variables with a histogram. 

---
## Question 3a

Before we produce any histograms using code, let's make sure we understand *how* to interpret the data plotted by a histogram. Run the following cell to create a histogram that shows the **distribution of penguin body mass (in grams)** for all the penguins in the `penguins` table.

```python
penguins.hist("body_mass_g", bins=np.arange(2000,7001, 500), density=False)
```

Based on the information presented in the above histogram, assign `num_penguins` to the **number of penguins** that have a body mass **between 3000 grams (*inclusive*) and 3500 grams (*exclusive*)**.  *Read this number from the histogram*

*Note*: Remember, in a histogram, each bin is inclusive of the left endpoint and exclusive of the right endpoint. What this means is that, for example, the bin between 2500 and 3000 grams above corresponds to body masses that are greater than or equal to 2500 grams and less than 3000 grams.

```python
num_penguins = ...
num_penguins
```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->

---
## Question 3b

Now it's time to create a histogram! For convenience, we've provided the `penguins` table in the cell below. There are four columns containing numerical variables: `"bill_length_mm"`, `"bill_depth_mm"`, `"flitter_length_mm"`, and `"body_mass_g"`.

```python
penguins.show(5)
```

In the cell below, write code that produces a histogram that visualizes the distribution of our penguins' flipper lengths. We've provided the bins for you to use as the `my_bins` variable. More specifically, you'll want to re-create the following histogram exactly:

<img src="images/examples/flipper_lengths.png">

```python

my_bins = np.arange(170,240, 5)
...

# Don't change any code below this line
plt.title("Distribution of Flipper Lengths (mm)");
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: 3)
```
>>> isinstance(sfo, Table)
True
```
```
>>> sfo.num_rows == 22869
True
```
```
>>> sfo.num_columns == 7
True
```


Tests for q1b (points: [1, 2])
```
>>> type(num_passengers) in (int, float, np.int64, np.float64)
True
```
```
>>> num_passengers == 14077961
True
```


Tests for q1c (points: [1, 1, 2])
```
>>> type(prop_europe) in (int, float, np.int64, np.float64)
True
```
```
>>> 0 <= prop_europe <= 1
True
```
```
>>> np.isclose(prop_europe, 0.08931552587646077)
True
```


Tests for q1d (points: 4)
```
>>> type(july_2019_mean) in (int, float, np.int64, np.float64)
True
```
```
>>> type(july_2020_mean) in (int, float, np.int64, np.float64)
True
```
```
>>> np.isclose(july_2019_mean, 75222.43333333333)
True
```
```
>>> np.isclose(july_2020_mean, 11177.619047619048)
True
```


Tests for q3a (points: 1)
```
>>> any(isinstance(num_penguins, x) for x in (int, float))
True
```
```
>>> 58 <= num_penguins <= 62
True
```
