```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10.worksheet.ipynb")
```

# Lab 10 - Worksheet 

Complete and turn in by Nov 2, 7pm. 



```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *
```

## WNBA Data 

We can use this dataset of statistics from the 2020 WNBA season to get more practice working with tables

```python
# Read in the data 

wnba = Table.read_table('data/wnba-2020.csv')
wnba
```

## Question 1 

Determine the number of rows and columns in the Table. 

```python
wnba_columns = ...
wnba_rows = ...
```

```python
grader.check("q1")
```

## Question 2 

Create a table of only the following columns: 'Player', 'Tm', 'Pos', 'G', 'PTS'

Name this table `wnba_pts` 

```python
# Create a table of only the following columns:
# 'Player', 'Tm', 'Pos', 'G', 'PTS'
# Name this table `wnba_pts`
wnba_pts = ...
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

## Question 3 

Show how to calculated the number of points scored per game for the `wnba_pts` table. 

```python
# Compute the number of points scored per game (Points Per Game)
...
```

<!-- END QUESTION -->

## Question 4

Add a Points Per Game column to the existing `wnba_pts` table.  Call this column `PPG`. 

```python
# Add a Points Per Game column to the existing `wnba_pts` table
...
wnba_pts.show(10)
```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 4)
```
>>> wnba_columns == 27
True
```
```
>>> wnba_rows == 162
True
```


Tests for q2 (points: 6)
```
>>> wnba_pts.num_rows == 162
True
```
```
>>> wnba_pts.num_columns == 5
True
```
```
>>> sorted(wnba_pts.labels) == ['G', 'PTS', 'Player', 'Pos', 'Tm']
True
```


Tests for q4 (points: 10)
```
>>> wnba_pts.num_columns == 6 
True
```
```
>>> sorted(wnba_pts.labels) == ['G', 'PPG', 'PTS', 'Player', 'Pos', 'Tm']
True
```
